# Parcourir une liste

Une des activités les plus fréquemment exécutées avec les listes c'est le parcours de la liste du premier au dernier élément. On trouve essentiellement deux méthodes pour parcourir l'intégralité d'une liste.

Première méthode : appliquer une boucle **for** sur un combinaison des fonction **range** et **len**.

```python
b = [2, 3, 5, 7, 11]

# len(b) donne le nombre d'élément de la liste b
print("Nombre d'éléments dans la liste :", len(b))

# range(n) permet de parcourir les entiers de 0 à n-1
print("On veut parcourir les entiers de", range(len(b)))

# Avec une boucle for on peut donc parcourir tous les rangs possibles avec la liste b.
print("Parcours de ces entiers un par un :")
for i in range(len(b)):
    print(i)
```

```python
b = [2, 3, 5, 7, 11]

# Pour parcourir les objets de la liste il suffit alors d'aller chercher chaque élément 
# avec la notation entre crochets du rang.

for i in range(len(b)):
    print("Element de rang", i, "->", b[i])
```

Seconde methode : appliquer une boucle **for** sur la liste elle-même. 

```python
c = ["Alpha", "Bravo", "Charlie", "Delta"]

print("Affichage des éléments de la liste un par un :")
for e in c:
    print("Element :", e)
```

Notez bien la différence entre les deux méthodes.

Dans le premier cas, la variable *i* parcourt la liste des rangs des objets de la liste alors que dans le second cas la variable _e_ parcourt directement la liste des objets.

La seconde méthode est plus courte, plus claire et plus rapide que la première mais elle a un défaut : on ne connaît pas le rang de l'objet qui est dans la variable *e*.

On choisira entre les deux méthodes suivant l'objectif qu'on souhaite atteindre.

A vous de jouer avec les trois exercices suivants :

<!-- grade_id: cell-0eb528a30bf5af53, points: 0, grade: true, solution: true -->
```python
# Exercice 1 :
# Faire unscript qui affiche chacun des entiers d'une liste en précisant s'il est pair ou impair
# Affichage attendu :
# 127 est impair
# 32 est pair
# 14 est pair
# 237 est impair
# 19 est impair
# 48 est pair

desEntiers = [127, 32, 14, 237, 19, 48]

### BEGIN SOLUTION
for e in desEntiers:
    if e % 2 == 0:
        print(f"{e} est pair")
    else:
        print(f"{e} est impair")   
### END SOLUTION
```

<!-- grade_id: cell-d73b9c2deedcbbde, points: 0, grade: true, solution: true -->
```python
# Exercice 2 :
# Faire un script qui pour chaque jour de la semaine affiche le nom du jour et son numéro 
# (de 1 pour lundi à 7 pour dimanche)
# Affichage attendu :
# Jour 1 : Lundi
# Jour 2 : Mardi
# Jour 3 : Mercredi
# Jour 4 : Jeudi
# Jour 5 : Vendredi
# Jour 6 : Samedi
# Jour 7 : Dimanche

semaine = ["Lundi", "Mardi", "Mercredi", "Jeudi", "Vendredi", "Samedi", "Dimanche"]

### BEGIN SOLUTION
for i in range(len(semaine)):
    print("Jour", i + 1, ":", semaine[i])
### END SOLUTION
```

<!-- grade_id: cell-b94121e2b935f291, points: 0, grade: false, solution: true -->
```python
# Exercice 3 :
# Faire une fonction qui reçoit une liste et un objet 
# et qui renvoie le rang de l'objet s'il est dans la liste, ou -1 s'iln'est pas dans la liste
# Il faut utiliser une boucle.

def trouverRang(liste, objet):
    ### BEGIN SOLUTION
    for i in range(len(liste)):
        if objet == liste[i]:
            return i
    return -1
    ### END SOLUTION
```

<!-- grade_id: cell-93e358fd22db2eb1, points: 0, grade: true, solution: false -->
```python
# Test de la fonction précédente
assert trouverRang(["S", "T", "O", "P"], "O") == 2
assert trouverRang([3, 6, 9, 12, 15, 18, 21], 21) == 6
assert trouverRang([1.1, 2.2, 3.3], 4.4) ==  -1
print("Les 3 tests sont passés avec succès.")
### BEGIN HIDDEN TESTS
assert trouverRang(list(range(10)), 5) == 5
assert trouverRang(list(range(8)), 20) == -1
assert trouverRang([], 11) == -1
### END HIDDEN TESTS
```